---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Testing do something crazy

```yaml
type: SingleProcessExercise
key: 2540e4d31d
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Testing 1

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Testing 3

```yaml
type: NormalExercise
key: 4a0d2b8ee6
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 4a20e4a081
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
test:
  1 2 3
```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
